# Assignment 1

For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web.

Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

```python
example_word_count()
```

Output:
```
5
```

## Part A

Find a list of all of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    # YOUR CODE HERE
    #raise NotImplementedError()
    search = "[A-Z][a-z]*"
    return re.findall(search, simple_string)
    

```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

```

```python
names()
```

Output:
```
['Amy', 'Mary', 'Ruth', 'Peter']
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    # YOUR CODE HERE
    # raise NotImplementedError()
    pattern = "[\w]*[\s][\w]*(?=:\ B)"
    return re.findall(pattern, grades)
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:

* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    new_list = []
    pattern = "(?P<host>\d{1,100}[.]\d{1,100}[.]\d{1,100}[.]\d{1,100})(\s[-]\s)(?P<user_name>\w*)(\s\[)(?P<time>[\w]*/[\w]*/.*)(\]\s\")(?P<request>.*)"
    for item in re.finditer(pattern,logdata):
        # print(item.groupdict())
        new_list.append(item.groupdict())
    return new_list

logs()
```

Output:
```
[{'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1" 302 4622'},
 {'host': '197.109.77.178',
  'user_name': 'kertzmann3129',
  'time': '21/Jun/2019:15:45:25 -0700',
  'request': 'DELETE /virtual/solutions/target/web+services HTTP/2.0" 203 26554'},
 {'host': '156.127.178.177',
  'user_name': 'okuneva5222',
  'time': '21/Jun/2019:15:45:27 -0700',
  'request': 'DELETE /interactive/transparent/niches/revolutionize HTTP/1.1" 416 14701'},
 {'host': '100.32.205.59',
  'user_name': 'ortiz8891',
  'time': '21/Jun/2019:15:45:28 -0700',
  'request': 'PATCH /architectures HTTP/1.0" 204 6048'},
 {'host': '168.95.156.240',
  'user_name': 'stark2413',
  'time': '21/Jun/2019:15:45:31 -0700',
  'request': 'GET /engage HTTP/2.0" 201 9645'},
 {'host': '71.172.239.195',
  'user_name': 'dooley1853',
  'time': '21/Jun/2019:15:45:32 -0700',
  'request': 'PUT /cutting-edge HTTP/2.0" 406 24498'},
 {'host': '180.95.121.94',
  
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

```